---
id: 68cae5b538ff798bbd4da004
title: "Challenge 66: Placeholder"
challengeType: 28
dashedName: challenge-66
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function placeholder(arg) {

  return arg;
}
```

# --solutions--

```js
function placeholder(arg) {

  return arg;
}
```
