---
id: 68cae5b538ff798bbd4da002
title: "Challenge 64: Placeholder"
challengeType: 28
dashedName: challenge-64
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function placeholder(arg) {

  return arg;
}
```

# --solutions--

```js
function placeholder(arg) {

  return arg;
}
```
