---
id: 68cae5b538ff798bbd4da002
title: "Challenge 64: Placeholder"
challengeType: 29
dashedName: challenge-64
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def placeholder(arg):

    return arg
```

# --solutions--

```py
def placeholder(arg):

    return arg
```
