---
id: 69b1028d6e265413d0198a2f
title: "Challenge 236: Placeholder"
challengeType: 29
dashedName: challenge-236
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def placeholder(arg):

    return arg
```

# --solutions--

```py
def placeholder(arg):

    return arg
```
