---
id: 69b5b2be76ec8135a7fbe973
title: "Challenge 244: Placeholder"
challengeType: 29
dashedName: challenge-244
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def placeholder(arg):

    return arg
```

# --solutions--

```py
def placeholder(arg):

    return arg
```
